Reactive Oxygen Detoxification Contributes to Mycobacterium abscessus Antibiotic Survival

Author response:

The following is the authors’ response to the previous reviews

Public Reviews:

Reviewer #1 (Public review):

Summary:

Persistence is a phenomenon by which genetically susceptible cells are able to survive exposure to high concentrations of antibiotics. This is especially a major problem when treating infections caused by slow growing mycobacteria such as M. tuberculosis and M. abscessus. Studies on the mechanisms adopted by the persisting bacteria to survive and evade antibiotic killing can potentially lead to faster and more effective treatment strategies.

To address this, in this study, the authors have used a transposon mutagenesis based sequencing approach to identify the genetic determinants of antibiotic persistence in M. abscessus. To enrich for persisters they employed conditions, that have been reported previously to increase persister frequency - nutrient starvation, to facilitate genetic screening for this phenotype. M.abs transposon library was grown in nutrient rich or nutrient depleted conditions and exposed to TIG/LZD for 6 days, following which Tnseq was carried out to identify genes involved in spontaneous (nutrient rich) or starvationinduced conditions. About 60% of the persistence hits were required in both the conditions. Pathway analysis revealed enrichment for genes involved in detoxification of nitrosative, oxidative, DNA damage and proteostasis stress. The authors then decided to validate the findings by constructing deletions of 5 different targets (pafA, katG, recR, blaR, Mab_1456c) and tested the persistence phenotype of these strains. Rather surprisingly only 2 of the 5 hits (katG and pafA) exhibited a significant persistence defect when compared to wild type upon exposure to TIG/LZD and this was complemented using an integrative construct. The authors then investigated the specificity of delta-katG susceptibility against different antibiotic classes and demonstrated increased killing by rifabutin. The katG phenotype was shown to be mediated through the production of oxidative stress which was reverted when the bacterial cells were cultured under hypoxic conditions. Interestingly, when testing the role of katG in other clinical strains of Mab, the phenotype was observed only in one of the clinical strains demonstrating that there might be alternative anti-oxidative stress defense mechanisms operating in some clinical strains.

Strengths:

While the role of ROS in antibiotic mediated killing of mycobacterial cells have been studied to some extent, this paper presents some new findings with regards to genetic analysis of M. abscessus susceptibility, especially against clinically used antibiotics, which makes it useful. Also, the attempts to validate their observations in clinical isolates is appreciated.

Weaknesses:

Amongst the 5 shortlisted candidates from the screen, only 2 showed marginal phenotypes which limits the impact of the screening approach.

We appreciate the reviewer’s comments, but we note that 4 out of 5 genes displayed phenotypes concordant with findings of the Tn-Seq data, with katG and pafA, as well as MAB_1456c (during starvation only) and blaR (in rich media only) having decreased survival as shown in Figure 3A-D. We do agree that some of the phenotypes were more modest in a single-mutant context than in the pooled Tn-Seq screen. In addition, several mutants that had modest changes in survival also showed profound defects in resuming growth after removal of antibiotics, with the pafA mutants particularly impaired. (Figure 3 - figure supplement 1).

While the role of KatG mediated detoxification of ROS and involvement of ROS in antibiotic killing was well demonstrated, the lack of replication of this phenotype in some of the clinical isolates limits the significance of these findings.

While the role of katG varied among strains, the antibiotic-induced accumulation of ROS was seen in all three strains (Figure 6A). This suggests that in some strains other ROS-detoxification pathways are able to compensate for the loss of katG.

(Figure 2—figure supplements 1–3)

Figure 1—figure supplement 1.

Reviewer #2 (Public review):

Summary:

The work set out to better understand the phenomenon of antibiotic persistence in mycobacteria. Three new observations are made using the pathogenic Mycobacterium abscessus as an experimental system: phenotypic tolerance involves suppression of ROS, protein synthesis inhibitors can be lethal for this bacterium, and levofloxacin lethality is unaffected by deletion of catalase, suggesting that this quinolone does not kill via ROS.

Strengths:

The ROS experiments are supported in three ways: measurement of ROS by a fluorescent probe, deletion of catalase increases lethality of selected antibiotics, and a hypoxia model suppresses antibiotic lethality. A variety of antibiotics are examined, and transposon mutagenesis identifies several genes involved in phenotypic tolerance, including one that encodes catalase. The methods are adequate for making these statements.

Weaknesses:

The work can be improved by a more comprehensive treatment of prior work, especially comparison of E. coli work with mycobacterial studies.

Moreover, the work still has some technical issues to fix regarding description of the methods, supplementary material, and reference formating.

See detailed responses below.

Overall impact: Showing that ROS accumulation is suppressed during phenotypic tolerance, while expected, adds to the examples of the protective effects of low ROS levels. Moreover, the work, along with a few others, extends the idea of antibiotic involvement with ROS to mycobacteria. These are fieldsolidifying observations.

Comments on revisions:

The authors have moved this paper along nicely. I have a few general thoughts.

It would be helpful to have more references to specific figures and panels listed in the text to make reading easier.

Text modified to add more figure references.

(1) I would suggest adding a statement about the importance of the work. From my perspective, the work shows the general nature of many statements derived from work with E. coli. This is important. The abstract says this overall, but a final sentence in the abstract would make it clear to all readers.

We appreciate the suggestion and have added a line to the abstract.

(2) The paper describes properties that may be peculiar to mycobacteria. If the authors agree, I would suggest some stress on the differences from E. coli. Also, I would place more stress on novel findings. This might be done in a section called Concluding Remarks. The paper by Shee 2022 AAC could be helpful in phrasing general properties.

We have added mention of this in the discussion (lines 354-356).

(3) Several aspects still need work to be of publication quality. Examples are the materials table and the presentation of supplementary material. Reference formatting also needs attention.

We respond to the specific details below.

Reviewer #3 (Public review):

Summary:

The manuscript demonstrates that starvation induces persister formation in M. abscesses.

They also utilized Tn-Seq for the identification of genes involved in persistence. They identified the role of catalase-peroxidase KatG in preventing death from translation inhibitors Tigecycline and Linezolid. They further demonstrated that a combination of these translation inhibitors leads to the generation of ROS in PBS-starved cells.

Strengths:

The authors used high-throughput genomics-based methods for identification of genes playing a role in persistence.

Weaknesses:

The findings could not be validated in clinical strains.

Comments on revisions: No more comments for the authors.

Recommendations for the authors:

Reviewer #1 (Recommendations for the authors):

The authors are strongly encouraged to check the references. There is some systematic error in the citations of references. Started to list but then they were too many.

For example Ln 51, Ref #11 cited, should be #10. Ln 59, #18 is wrongly cited. Should be - Ln 104. Ref #27 wrongly cited.

Ref #26 and #28 identical.

Even in discussion section a lot of references are mis-cited.

We very much appreciate the reviewer catching this issue with the import of our references and we have corrected this.

Reviewer #2 (Recommendations for the authors):

Below I have listed comments on specific issues that I hope are useful during revision.

Line 21 population is singular

Text modified

Line 21 comma after antibiotic (subordinate clause) Line

Text modified

25 is how singular?

Text modified

Impression of abstract: the work seems to confirm and therefore generalize concepts derived from studies with E. coli. If the authors agree, such a statement would be appropriate as a final sentence. I would also look for novel features to stress in the abstract.

Line 41 this challenge is vague

Text modified

Line 43 comma such as (also comma at the end of the parenthetical statement). This type of comma error is common throughout the manuscript and slows reading.

Text modified

Line 60 paradoxically. Is this the best concept? Or is it the natural effect of evolution (assuming that mycobacteria or their ancestors were exposed to environmental antibiotics)?

It is certainly problematic for clearing infection.

Text not modified.

Line 63 highlighted uncertainties ... meaning is unclear especially since you may have changed what "model" is referring to.

Text modified

Line 66 models.... Do you really mean systems? Models of what?

This refers to mechanistic models. Text not modified.

Line 67 arrest cell division. This is written as if it were true. Does the evidence point specifically to cell division or perhaps more accurately suppression of metabolism (see Ye et al 2025 mBio).

Both have been postulated as important. Text modified to add concept of metabolism

... targeted by antibiotics non-essential... Do you think that antibiotics work by inactivating essential targets? That seems overly simplistic, as lethal action is more likely the metabolic response to the damage caused. By the end of the paragraph you come around to this view, but you have already misdirected the reader. The reader is not sure what to believe. Line 70 note that there are many inhibitors of transcription and translation that only block growth, they do not rapidly kill cells

There can be both direct, and indirect secondary killing mechanisms. We devote a significant portion of the Discussion section to this topic.

Line 71 debate. There was indeed a debate, but reference 22 is not a valid citation for this. I think you mislead the reader by not accurately describing the debate. It was basically about the inability of Kim Lewis and James Imlay to reproduce the work of ref. 22. A great deal of prior work and then subsequent work showed that the challenge to ref. 22 lacked substance.

(1) Text modified to fix an error in the citation number related to direct β-lactam-mediated lysis.

(2) We agree that there is a great deal of data supporting antibiotic-induced ROS as important for bactericidal activity in many circumstances and do not argue otherwise. This sentence points out that over the years the paradigm for how antibiotics kill bacteria has evolved.

Line 80. It seems you are starting a new topic here. What about beginning a new paragraph?

The paragraph introduces mycobacteria of which Mabs is one. Text not modified.

Line 85 delete the comma: it implies a compound sentence that is not delivered.

Text modified.

Line 109 screen singular

Text modified.

Line 156 these conditions is imprecise and vague

Conditions were described in paragraph above in the manuscript. Text not modified.

Fig 2 it would be helpful to more clearly define the meaning of the coordinates

Text modified.

Line 230 and throughout please indicate the location of the data being cited for rapid reader reference

Text modified.

Lines 315-323 You could use this paragraph as the first of the Discussion. Some readers prefer to read the Discussion before the results. For them, a summary at the beginning of the Discussion is useful.

Text modified.

Line 328 without underlying mechanism... for E. coli refer to Zeng PNAS 2022. Depending on when the final version of this paper happens, there should be a figure in a Zhao Zhu mLife paper on purA that will have been published. Since it is not yet available, it cannot be cited.

We agree that the Zeng et al study is interesting and have added this reference to our discussion. However, these findings related to broad Crp-regulated tolerance actually underscore the point that we are making: that there are multiple factors (Crp, RelA, Lon, TisB, MazE, others) that mediate antibiotic tolerance.

Line 339 where are the data?

These data are in Figure 5, panels C, D. We have clarified the text to indicate that only a single agent from each of these classes was tested.

Line 346 here you are summarizing evidence for ROS in killing mycobacteria. You should include the moxifloxacin study by Shee et al 2022 AAC.

Reference added.

Line 348 refer to James Collins' work with E. coli in which his lab examined agents with a variety of mechanisms. There seems to be a fundamental difference between E. coli and mycobacteria with respect to rifampicin, a strictly static agent in E. coli but clearly lethal in mycobacteria. Note that chloramphenicol is static in E. coli and blocks ROS production. What does it do in mycobacteria? A brief discussion of this difference might be relevant at line 362

Text modified.

Lines 364-368 Here the idea might be simply that there are two modes of killing, one that is a direct extension of class-specific damage (chromosome fragmentation with fluoroquinolones, for example, or cell lysis by beta-lactams) and a second that is a metabolic response to the antibiotic damage (ROS accumulation). The second type is not class specific. Within this context, the mycobacterial killing by rifampicin might be a class-specific extension of inhibition of transcription that does not occur in E. coli.

Agreed, text modified to include this.

Line 400 The Key Resource table is not of publication quality. Precision and repeatability can be improved by spelling out the name of the vendor and its location (City, Country). In the present case, use of BD is lab jargon.

We appreciate the reviewer’s precision. However, this is actually not lab jargon. Becton, Dickinson and Company now refers to itself as BD (see https://www.bd.com/en-us), and the American Type Culture Collection now refers to itself as ATCC (see https://www.atcc.org/about-us/who-we-are).

Line 639 It would be good to have experienced colleagues critically review the manuscript, especially for English usage. Listing those persons here adds to the credibility of the work

Text not changed.

References: please refer to the journal style. Here you use italic for titles and scientific names, thereby obscuring the scientific names. Normally article titles are not italic and scientific names are ALWAYS italic unless prohibited by journal style.

Our reference format is concordant with eLife submission guidelines, and all references are reformatted by the journal at the time of final publication (see https://elifesciences.org/insideelife/a43f95ca/elife-references-yes-we-take-any-format-no-we-re-not-rekeying).

Supplemental Material: Please refer to journal style. Normally this is a stand-alone document that includes a title page and carefully crafted figure legends. Supplemental figures would be numbered as 1, 2, ... A professional appearing Supplemental Material section shows author publication experience not obvious in other parts of the paper. The text indicated MIC determinations. I would like to see a table of MIC values.

(1) MIC table added as Supplemental Table 5.

(2) The Supplemental figures are submitted and named in accordance with eLife instructions. Please note that for eLife, there is not a stand-alone supplementary figure section with a title page as you are requesting, but instead the figure supplements for each figure are provided as online files linked to each figure.

Changes from the earlier version to the revised version

Abstract

Earlier version: When a population of bacteria are exposed to a bactericidal antibiotic most cells die rapidly. However, a sub-population of antibiotic-tolerant cells known as “persister cells” can survive for prolonged periods, and antibiotic tolerance can be strongly induced in a population by stresses such as nutrient deprivation. However, the pathways required to maintain viability in this setting, and how they are regulated are poorly understood. To identify genetic determinants of antibiotic tolerance in mycobacteria, we carried out transposon mutagenesis insertion sequencing (Tn-Seq) screens in Mycobacterium abscessus (Mabs) exposed to bactericidal translation-inhibiting antibiotics. This analysis identified genes essential for the survival of both spontaneous persister cells as well as for stress-induced tolerance, allowing the first genetic comparison of these states in mycobacteria. Pathway analysis identified multiple genes involved in the detoxification of reactive oxygen species (ROS), including the catalase-peroxidase katG, which contributed to survival in both stressed and un-stressed populations. In addition, we found that endogenous ROS were generated following antibiotic exposure, and that hypoxia significantly impaired bacterial killing. Thus, in Mabs, the lethality of some antibiotics is amplified by toxic ROS accumulation, and antibiotic-tolerant cells require detoxification systems to remain viable.

Revised version: When a population of bacteria is exposed to a bactericidal antibiotic, most cells die rapidly. However, a sub-population of antibiotic-tolerant cells known as “persister cells” can survive for prolonged periods. In addition, antibiotic tolerance can be broadly induced throughout the population by stresses such as nutrient deprivation. However, the pathways required to maintain viability in this setting, and how stress induces antibiotic tolerance are both poorly understood. To identify genetic determinants of antibiotic tolerance in mycobacteria, we carried out transposon mutagenesis insertion sequencing (Tn-Seq) screens in Mycobacterium abscessus (Mabs) exposed to bactericidal translation-inhibiting antibiotics. This analysis identified genes essential for the survival of both spontaneous persister cells as well as for stress-induced tolerance, allowing the first genetic comparison of these states in mycobacteria. Pathway analysis identified multiple genes involved in the detoxification of reactive oxygen species (ROS), including the catalase-peroxidase katG, which contributed to survival in both unstressed and nutrient-starved cells. In addition, we found that endogenous ROS were generated by translation-inhibiting antibiotics, and that hypoxia impaired bacterial killing. KatG specifically contributed to survival following exposure to transcription or translation inhibitors, but not other antibiotic classes tested. Thus, the lethality of some antibiotics is amplified by toxic ROS accumulation, and antibiotic-tolerant cells require detoxification systems to remain viable. These findings further demonstrate that antibiotic-induced ROS plays a broad role in mediating antibiotic lethality across diverse organisms.

Introduction

Earlier version: A key goal of antibiotic therapy is, in conjunction with the immune system, to eradicate the infecting bacteria. While many common bacterial infections respond rapidly to antibiotics, with 1-2 weeks of therapy sufficient to achieve high cure rates, there are also infections where bacterial clearance is slow and often incomplete1, 2. This challenge is exemplified by mycobacterial infections. Fully-susceptible Mtb requires multiple antibiotics for four months or longer3, 4, and infections by non-tuberculous mycobacteria (NTM) such as Mycobacterium abscessus (Mabs) are often even more difficult to eradicate; Mabs often requires treatment for 12-18 months, and even then has a 50% relapse rate5, 6.

Revised version: A key goal of antibiotic therapy is, in conjunction with the immune system, to eradicate the infecting bacteria. While many common bacterial infections respond rapidly to antibiotics, with 1-2 weeks of therapy sufficient to achieve high cure rates, there are also infections where bacterial clearance is slow and often incomplete1, 2. This challenge is exemplified by mycobacterial infections. Fully-susceptible Mtb requires multiple antibiotics for four months or longer3, and infections by non-tuberculous mycobacteria (NTM), such as Mycobacterium abscessus (Mabs), are even more difficult to eradicate; Mabs often requires treatment for 12-18 months, and even then has a 50% relapse rate4, 5.

Earlier version: While the ability of mycobacteria to escape antibiotic-mediated killing is multifactorial, the phenomenon of antibiotic tolerance is likely an important contributor7–10. Studies dating from the 1940s noted that when a population of susceptible bacteria were exposed to a bactericidal antibiotic such as penicillin, the majority of the population died within a few hours, but that a small sub-population of cells remained viable for days11. Importantly, these antibiotic-tolerant cells referred to as “persister cells” had not acquired a mutation conferring heritable antibiotic resistance, and do not grow in the presence of the antibiotic. Rather, they had entered into a readily-reversable phenotypic state where, despite antibiotic-mediated inhibition of critical processes, they are able to survive12–14. In addition, a number of physiologic stresses increase antibiotic tolerance in a population, as bacterial cell death is markedly slowed by stresses such as nutrient deprivation or acidic pH10, 11, 15, 16. Notably, these same stresses are encountered in the lysosome of an activated immune cell17, and studies of pathogens isolated from activated macrophages indeed show a strong immune-mediated increase antibiotic tolerance9, 18. Thus, paradoxically, the immune system can actually impede bacterial eradication by antibiotics.

Revised version: While the ability of mycobacteria to escape antibiotic-mediated killing is multifactorial, the phenomenon of antibiotic tolerance is likely an important contributor6–9. Studies dating from the 1940s noted that when a population of susceptible bacteria were exposed to a bactericidal antibiotic such as penicillin, the majority of the population died within a few hours, but that a small sub-population of cells remained viable for days10. Importantly, these antibiotic-tolerant cells referred to as “persister cells” had not acquired a mutation conferring heritable antibiotic resistance, and do not grow in the presence of the antibiotic. Rather, they had entered into a readily reversable phenotypic state where, despite antibiotic-mediated inhibition of critical processes, they were able to survive11–13. In addition, a number of physiologic stresses increase antibiotic tolerance in a population, as bacterial cell death is markedly slowed by stresses such as nutrient deprivation or acidic pH9, 10, 14, 15. Notably, these same stresses are encountered in the lysosome of an activated immune cell16, and studies of pathogens isolated from activated macrophages indeed show a strong immune-mediated increase antibiotic tolerance8, 17. Thus, paradoxically, the immune system may actually impede bacterial eradication by antibiotics.

Earlier version: Antibiotic tolerance has been studied extensively in model systems such as Escherichia coli, which has provided important insights, but also highlighted uncertainties of current models. Several different regulators of antibiotic tolerance have been identified in E. coli including the HipBA toxin-antitoxin system19, guanosine pentaphosphate ((p)ppGpp) synthesis by RelA/SpoT enzymes20, 21, and the activity of Lon protease20. In each of these models, the postulated mechanism is to arrest cell division and render the process targeted by antibiotics non-essential. However, important questions remain unanswered. It is unclear how cells remain viable when critical processes such as transcription or translation are blocked by antibiotics, as well as how the process is regulated and induced by stress. Even the mechanism of cell death following antibiotic exposure itself has been a matter of debate - originally antibiotics were presumed to kill bacteria as a direct result of inhibition of their target molecule, such as β-lactam antibiotics disrupting cell wall integrity, directly leading to mechanical cell lysis22. However, a number of studies, largely from E. coli, have demonstrated that in addition to the initial target inhibition, which may be bacteriostatic, bactericidal antibiotics can also cause secondary lethal ROS accumulation, leading to cell death14, 23–28.

Revised version: Antibiotic tolerance has been studied extensively in model systems such as Escherichia coli, which has provided important insights, but also highlighted uncertainties of current paradigms. Several different regulators of antibiotic tolerance have been identified in E. coli including the HipBA toxin-antitoxin system18, guanosine pentaphosphate ((p)ppGpp) synthesis by RelA/SpoT enzymes 19–22, and Lon protease23. In each of these models, the postulated mechanism is to slow metabolism or cell division and render the process targeted by antibiotics non-essential. However, important questions remain. It is unclear how cells remain viable when critical processes, such as transcription or translation, are blocked by antibiotics, as well as how antibiotic tolerance is induced by stress. Even the mechanism of cell death following antibiotic exposure itself has been a matter of debate - originally antibiotics were presumed to kill bacteria as a direct result of inhibition of their target molecule, such as β-lactam antibiotics disrupting cell wall integrity, directly leading to mechanical cell lysis24. However, a number of studies, largely from E. coli, have demonstrated that in addition to the initial target inhibition, which may be bacteriostatic, bactericidal antibiotics can also cause secondary lethal ROS accumulation, leading to cell death13,25–30.

Earlier version: Mycobacterial persister cells are particularly resilient, as Mycobacterium smegmatis (Msmeg) and Mycobacterium tuberculosis (Mtb) persisters can endure many weeks of antibiotic exposure, and stress-induced antibiotic tolerance develops readily10, 14, 29. Mabs is a species of rapidly-growing mycobacteria, with a doubling time of ∼4h in rich media30, and while often environmental, it causes opportunistic infections in patients with structural lung disease such as cystic fibrosis. It is also among the most difficult of all bacterial pathogens to treat, because in addition to forming persister cells, it is also intrinsically resistant to many classes of antibiotics, leaving few treatment options5. This leads to the use of antibiotics with greater toxicity to patients, and a need to use these agents for prolonged periods to prevent relapse. Thus, identifying the genes that Mabs persister cells rely on for survival, as well as the genes involved in the induction and maintenance of stress-induced antibiotic tolerance might highlight pathways that could be targeted therapeutically to eliminate antibiotic-tolerant cells.

Revised version: Mycobacterial persister cells are particularly resilient, as Mycobacterium smegmatis (Msmeg) and Mycobacterium tuberculosis (Mtb) persisters can endure many weeks of antibiotic exposure, and stress-induced antibiotic tolerance develops readily9, 13, 28. Mabs is a species of rapidly-growing mycobacteria, with a doubling time of ∼4h in rich media31, and while often environmental, it causes opportunistic infections in patients with structural lung disease such as cystic fibrosis. It is also among the most difficult of all bacterial pathogens to treat, because in addition to forming persister cells, it is also intrinsically resistant to many classes of antibiotics, leaving few treatment options5. This leads to the use of antibiotics with greater toxicity to patients and a need to use these agents for prolonged periods to prevent relapse. Thus, identifying the genes that Mabs persister cells rely on for survival, as well as the genes involved in the induction and maintenance of stress-induced antibiotic tolerance might highlight pathways that could be targeted therapeutically to eliminate antibiotic-tolerant cells.

Earlier version: Previous genetic screens have studied antibiotic responses in mycobacteria, with some evaluating heritable resistance, and others investigating tolerance. Several studies of resistance have successfully used either transposon mutagenesis with insertion site sequencing (Tn-Seq) or CRISPR-based transcriptional repression with high-throughput sequencing of guide RNAs (CRISPRi) to identify genes promoting growth in sub-inhibitory concentrations of antibiotic. These studies have provided insights such as highlighting the importance of cell membrane permeability as a mechanism controlling antibiotic penetration into the cytoplasm31–33.

Revised version: Previous genetic screens have studied antibiotic responses in mycobacteria, with some evaluating heritable resistance, and others investigating tolerance. Several studies of resistance have successfully used either transposon mutagenesis with insertion site sequencing (Tn-Seq) or CRISPR-based transcriptional repression with high-throughput sequencing of guide RNAs (CRISPRi) to identify genes promoting growth in sub-inhibitory concentrations of antibiotic. These studies have provided insights, such as highlighting the importance of cell membrane permeability as a mechanism controlling antibiotic penetration into the cytoplasm32–34.

Earlier version: Persister formation has proven challenging to study, likely because the low frequency of persisters leads to population bottlenecks that confound genetic analysis. Although screens in Mtb have been conducted in macrophages and mice, and genes such as glpK and cinA identified, overall the number of mutants isolated in these screens has been low34–35. There has been one effective in vitro Tn-Seq study examining antibiotic tolerance in Mtb exposed to starvation and rifampin that isolated over 100 mutants27, demonstrating the feasibility of genetic screening in this context. However, whether these phenotypes seen with rifampin in Mtb extend to other mycobacteria and other antibiotics remain to be determined.

Revised version: Persister formation has proven challenging to study, likely because the low frequency of persisters leads to population bottlenecks that confound genetic analysis. Although screens in Mtb have been conducted in macrophages and mice, and genes, such as glpK and cinA identified, overall the number of mutants isolated in these screens has been low35, 36. There has been one effective in vitro Tn-Seq study examining antibiotic tolerance in Mtb exposed to starvation and rifampin that isolated over 100 mutants29, demonstrating the feasibility of genetic screening in this context. However, whether these phenotypes seen with rifampin in Mtb extend to other mycobacteria and other antibiotics remains to be determined.

Earlier version: Here, we study antibiotic tolerance in Mabs and describe the results of a genome-wide Tn-seq screens seeking to identify the genes required for both spontaneous persister cell survival as well as starvation-induced antibiotic tolerance following exposure to translation-inhibiting antibiotics. We identified several discrete processes contributing to survival and observed a prominent role for ROS detoxifying factors such as the catalase-peroxidase enzyme KatG, which contributed to both spontaneous persister survival and starvation-induced antibiotic tolerance. Consistent with the protection conferred by KatG, we found that endogenous ROS accumulated following antibiotic exposure, and that the removal of oxygen significantly impaired bacterial killing. Taken together these findings support a model in mycobacteria where the lethality of translation-inhibiting antibiotics is amplified by a secondary accumulation of toxic ROS, and that survival requires active detoxification systems.

Revised version: Here, we study antibiotic tolerance in Mabs and describe the results of genome-wide Tn-seq screens seeking to identify the genes required for both spontaneous persister cell survival as well as starvation-induced antibiotic tolerance following exposure to translation-inhibiting antibiotics. We identified several discrete processes contributing to survival and observed a prominent role for ROS detoxifying factors, such as the catalase-peroxidase enzyme KatG, which contributed to both spontaneous persister survival and starvation-induced antibiotic tolerance. Consistent with the protection conferred by KatG, we found that endogenous ROS accumulated following antibiotic exposure, and that the removal of oxygen significantly impaired bacterial killing. Taken together these findings support a model in mycobacteria where the lethality of translation-inhibiting antibiotics is amplified by a secondary accumulation of toxic ROS, and that survival requires active detoxification systems.

Results

Starvation-induced antibiotic tolerance in mycobacteria

Earlier version: We first sought to develop conditions suitable for genetic analysis of persister cell survival and stress-induced antibiotic tolerance mycobacteria. Genetic screens examining persister cell physiology face two inherent obstacles. First, these cells are rare in unstressed bacterial populations, and antibiotic-mediated cell death creates population bottlenecks that obscure mutant phenotypes. Second, most mycobacterial populations contain spontaneous drug-resistant mutants that can expand if the population is exposed to a single antibiotic. To overcome these obstacles, we sought to establish large-scale, high-density culture conditions to prevent genetic bottlenecks, and used multiple antibiotics to suppress spontaneous drug-resistant mutants. We began by assessing the feasibility of this approach using wild-type Msmeg. We exposed the cells to either the combination of rifampin, isoniazid, and ethambutol (RIF/INH/EMB) used to treat Mtb, or to the combination of tigecycline and linezolid (TIG/LZD), two translation-inhibiting antibiotics frequently used to treat Mabs6, 36, and empirically determined the minimum inhibitory concentrations (MICs) and minimum bactericidal concentrations (MBC) for each antibiotic under the high-density culture conditions that would be needed for genetic analysis of persister cells. Both antibiotic combinations reduced the bacterial population >1000-fold within 72h (Figure 1A). We then evaluated both spontaneous persister formation and stress-induced tolerance under these conditions in Msmeg. We compared logarithmically growing (mid-log) cultures in 7H9 rich media to cultures starved for 2 days in phosphate buffered saline (PBS) prior to addition of antibiotics. Consistent with expectations, we found a marked increase in antibiotic tolerance in starved cultures, with a 100-fold increase in survival following TIG/LZD exposure and a 10,000-fold increase following RIF/INH/EMB exposure (Figure 1A).

Revised version: We first sought to develop conditions suitable for genetic analysis of persister cell survival and stress-induced antibiotic tolerance mycobacteria. Genetic screens examining persister cell physiology face two inherent obstacles. First, these cells are rare in unstressed bacterial populations, and antibiotic-mediated cell death creates population bottlenecks that obscure mutant phenotypes. Second, most mycobacterial populations contain spontaneous drug-resistant mutants that can expand if the population is exposed to a single antibiotic. To overcome these obstacles, we sought to establish large-scale, high-density culture conditions to prevent genetic bottlenecks, and used multiple antibiotics to suppress spontaneous drug-resistant mutants. We began by assessing the feasibility of this approach using wild-type Msmeg. We exposed the cells to either the combination of rifampin, isoniazid, and ethambutol (RIF/INH/EMB) used to treat Mtb, or to the combination of tigecycline and linezolid (TIG/LZD), two translation-inhibiting antibiotics frequently used to treat Mabs5,37, and empirically determined the minimum inhibitory concentrations (MICs) and minimum bactericidal concentrations (MBC) for each antibiotic under the high-density culture conditions that would be needed for genetic analysis of persister cells. Both antibiotic combinations reduced the bacterial population >1000-fold within 72h (Figure 1A). We then evaluated both spontaneous persister formation and stress-induced tolerance under these conditions in Msmeg. We compared logarithmically growing (mid-log) cultures in 7H9 rich media to cultures starved for 2 days in phosphate buffered saline (PBS) prior to addition of antibiotics. Consistent with expectations, we found a marked increase in antibiotic tolerance in starved cultures, with a 100-fold increase in survival following TIG/LZD exposure and a 10,000-fold increase following RIF/INH/EMB exposure (Figure 1A).

Starvation induces antibiotic tolerance in diverse mycobacteria.

(A) Msmeg, (B) Mtb or (C) Mabs were grown in 7H9 rich media or starved in PBS prior to the addition of antibiotics and surviving colony forming units (CFU) enumerated. For the rapidly growing mycobacteria, Msmeg and Mabs, cells were allowed to adapt for 48h prior to antibiotics, for slow-growing Mtb, cells were allowed to adapt for 14-21d prior to antibiotics. Samples without pre-adaptation were washed and placed directly into PBS with antibiotics. Antibiotic concentrations were: Msmeg - Isoniazid (INH) 32μg/ml (8 x MIC), rifampin (RIF) 32 μg/ml (8 x MIC), ethambutol (EMB) 4μg/ml (8 x MIC), tigecycline (TIG) 1.25 μg/ml (8 x MIC), linezolid (LZD) 2.5 μg/ml (8 x MIC). Mtb – RIF 0.1 μg/ml (4 x MIC), INH at 0.1 μg/ml (4 x MIC), EMB at 8 μg/ml (4 x MIC). Mabs – TIG 10 µg/ml (8 x MIC), LZD 100 μg/ml (20 x MIC). Antibiotics with half-lives shorter than the duration of experiment were re-added at the following intervals: TIG, EMB every 3d; RIF, INH every 6d. Error bars represent SEM, statistical significance is calculated at each time point using student’s t test. ****: p<0.0001, ***: p<0.001, **: p<0.01, *: p<0.05, ns: p>0.05. Data are combined from 3 independent experiments.

Earlier version: We next examined two species of pathogenic mycobacteria under these conditions to similarly assess starvation-induced antibiotic tolerance37–40. We again compared cells starved in PBS to logarithmically growing cells in 7H9 and found that cultures of wild-type Mabs (ATCC 19977) and Mtb (Erdman) also displayed dramatic increases antibiotic tolerance in nutrient-deprived cultures (Figure 1B,C). Notably, for Mabs and Msmeg, the development of tolerance required an adaptation period of several days under starvation conditions, as survival was dramatically impaired if cells were shifted immediately into nutrient-deficient conditions with antibiotics, suggesting that a regulated process needed to be completed. Surprisingly, Mtb tolerance developed rapidly without pre-adaptation suggesting that this organism might have additional response pathways enabling more rapid adaptation.

Revised version: We next examined two species of pathogenic mycobacteria to assess starvation-induced antibiotic tolerance. Both Mabs and Mtb have been shown to display this response, and we sought to determine whether this could be observed under the conditions needed to conduct a Tn-Seq screen38–41. We again compared cells starved in PBS to logarithmically growing cells in 7H9 and found that, under these conditions, cultures of wild-type Mabs (ATCC 19977) and Mtb (Erdman) displayed dramatic increases antibiotic tolerance in nutrient-deprived cultures (Figure 1B,C). Notably, for Mabs and Msmeg, the development of tolerance required an adaptation period of several days under starvation conditions, as survival was dramatically impaired if cells were shifted immediately into nutrient-deficient conditions with antibiotics, suggesting that a regulated process needed to be completed. Surprisingly, Mtb tolerance developed rapidly without pre-adaptation suggesting that this organism might have additional response pathways enabling more rapid adaptation.

Identification of pathways needed for antibiotic tolerance in Mabs

Earlier version: We used these conditions to carry out Tn-Seq screens in Mabs to identify genes necessary for both the survival of spontaneous persister cells and for starvation-induced antibiotic tolerance. We conducted the screen using a Mabs Himar1 Tn library comprised of ∼55,000 mutations across ∼91,000 possible TA insertion sites covering all 4,992 non-essential Mabs genes in strain ATCC 1997733. To study spontaneous persister cells, cultures were maintained in continuous log-phase in 7H9 rich media for 48h prior to antibiotic exposure and then exposed to TIG/LZD for 6 days (Figure 2A) a point at which spontaneous persister cells comprise the majority of the population (Figure 1C). To study starvation-induced tolerance, cultures were starved in PBS for 48h prior to antibiotic exposure in PBS. Following antibiotic treatment, cells were then washed and resuspended in antibiotic-free liquid media to recover and passaged 1:100 three times in continuous log-phase to expand surviving cells. We then isolated genomic DNA, sequenced the Tn insertion sites, and used TRANSIT software41 to quantify the abundance of each Tn mutant across different conditions to identify genes with statistically-significant differences in distribution. We identified 277 Mabs genes required for surviving TIG/LZD exposure in rich media, 271 genes required for survival during starvation and 362 genes required to survive the combined exposure to antibiotics and starvation (Log2 fold-change > 0.5 and Benjamini–Hochberg adjusted p-value (p-adj.) ≤ 0.05; Figure 2B-E). Of the genes required for survival, ∼60% were required in both nutrient-replete and starvation states, although condition-specific determinants were also seen (Figure 2F). As expected, we identified genes with already-established functions in antibiotic responses, including MAB_2752 and MAB_2753 which are both homologs of known antibiotic transporters in Mtb, as well as tetracycline-responsive transcription factors like MAB_4687 and MAB_0314c (Table S1), indicating an ability of these Tn-Seq conditions to identify physiologically relevant genes known to mitigate antibiotic stress.

Revised version: We used these conditions to carry out Tn-Seq screens in Mabs to identify genes necessary for both the survival of spontaneous persister cells and for starvation-induced antibiotic tolerance. We conducted the screen using a Mabs Himar1 Tn library comprised of ∼55,000 mutations across ∼91,000 possible TA insertion sites covering all 4,992 non-essential Mabs genes in strain ATCC 1997734. To study spontaneous persister cells, cultures were maintained in continuous log-phase in 7H9 rich media for 48h prior to antibiotic exposure and then exposed to TIG/LZD for 6 days (Figure 2A) a point at which spontaneous persister cells comprise the majority of the population (Figure 1C). To study starvation-induced tolerance, cultures were starved in PBS for 48h prior to antibiotic exposure in PBS. Following antibiotic treatment, cells were then washed and resuspended in antibiotic-free liquid media to recover and passaged 1:100 three times in continuous log-phase to expand surviving cells. We then isolated genomic DNA, sequenced the Tn insertion sites, and used TRANSIT software42 to quantify the abundance of each Tn mutant across different conditions to identify genes with statistically-significant differences in distribution. We identified 277 Mabs genes required for surviving TIG/LZD exposure in rich media, 271 genes required for survival during starvation and 362 genes required to survive the combined exposure to antibiotics and starvation (Log2 fold-change > 0.5 and Benjamini–Hochberg adjusted p-value (p-adj.) ≤ 0.05; Figure 2B-E). Of the genes required for survival, ∼60% were required in both nutrient-replete and starvation states, although condition-specific determinants were also seen (Figure 2F). As expected, we identified genes with already-established functions in antibiotic responses, including MAB_2752 and MAB_2753 which are both homologs of known antibiotic transporters in Mtb, as well as tetracycline-responsive transcription factors like MAB_4687 and MAB_0314c (Supplemental Table 1), indicating an ability of these Tn-Seq conditions to identify physiologically relevant genes known to mitigate antibiotic stress.

Tn-Seq identifies genes required for antibiotic tolerance in Mabs.

Earlier version: (A) Experimental design. (B-E) Tn-Seq analysis showing relative abundance of individual genes under the indicted conditions. Genes depleted relative to the input population have negative values. All cultures were fully aerated throughout the experiment and cultures without antibiotics received and equal volume of DMSO. For (B-D) gene abundance in the indicated condition is measured relative the input log-phase population. In (E) an additional comparison is made for PBS with antibiotics relative to the PBS condition. Genes with significant decreases in abundance are shown in color (p-adj. < 0.05 and log2 fold-change > 0.5) using the Benjamini–Hochberg adjustment for multiple hypothesis testing. (F) Number of genes essential in each condition relative to the input population. (G) Pathway enrichment analysis of the essential genes in each condition using the DAVID knowledgebase (p <0.05). Screens were run as 3 independent experiments and the combined results analyzed. Antibiotic conditions were as described above.

Revised version: (A) Experimental design. (B-E) Tn-Seq analysis showing relative abundance of individual genes under the indicted conditions. For (B-D) gene abundance in each condition is measured relative the input with negative values for genes depleted in each experimental condition relative to the input. Log2 fold-change is on x-axis with -log10 of the p-value on y-axis. All cultures were fully aerated throughout the experiment and cultures without antibiotics received and equal volume of DMSO. In (E) an additional comparison is made for PBS with antibiotics relative to the PBS condition. Genes with significant decreases in abundance are shown in color (p-adj. < 0.05 and log2 fold-change > 0.5) using the Benjamini–Hochberg adjustment for multiple hypothesis testing. (F) Number of genes essential in each condition relative to the input population. (G) Pathway enrichment analysis of the essential genes in each condition using the DAVID knowledgebase (p <0.05). Screens were run as 3 independent experiments and the combined results analyzed. Antibiotic conditions were as described above.

Earlier version: To identify other cellular processes necessary for survival we performed pathway enrichment analysis on the set of genes identified by Tn-Seq. We used the DAVID42 analysis tool to perform systematic queries of the KEGG, GO, and Uniprot databases and identify over-represented processes and pathways. Interestingly, although cells were exposed to translation-inhibiting antibiotics, and no exogenous oxidative or nitrosative stress was applied, we identified a number of factors needed to combat these stresses in spontaneous persister cells. This included bfrB (bactoferritin), ahpE (peroxiredoxin) and katG (catalase/peroxidase) as well as 5 components of the bacterial proteasome pathway, known to mediate resistance to nitrosative stress in Mtb43 (Figure 2G, Table S2). We also identified multiple members of DNA-damage response pathways including recF, recG, uvrA, uvrB and uvrC. Examining genes required for starvation-induced tolerance, a number of the same pathways were again seen, and the mutant with the greatest survival defect in this context was mntH, a redox-regulated Mn2+/Zn2+ transporter implicated in peroxide resistance in other organisms44–45.

Revised version: To identify other cellular processes necessary for survival we performed pathway enrichment analysis on the set of genes identified by Tn-Seq. We used the DAVID43 analysis tool to perform systematic queries of the KEGG, GO, and Uniprot databases and identify over-represented processes and pathways. Interestingly, although cells were exposed to translation-inhibiting antibiotics, and no exogenous oxidative or nitrosative stress was applied, we identified a number of factors needed to combat these stresses in spontaneous persister cells. This included bfrB (bactoferritin), ahpE (peroxiredoxin) and katG (catalase/peroxidase) as well as 5 components of the bacterial proteasome pathway, known to mediate resistance to nitrosative stress in Mtb44 (Figure 2G, Supplemental Table 2). We also identified multiple members of DNA-damage response pathways including recF, recG, uvrA, uvrB and uvrC. Examining genes required for starvation-induced tolerance, a number of the same pathways were again seen, and the mutant with the greatest survival defect in this context was mntH, a redox-regulated Mn2+/Zn2+ transporter implicated in peroxide resistance in other organisms45, 46.

Earlier version: To independently confirm a role in antibiotic tolerance for a set of genes from diverse pathways that were identified by Tn-Seq we selected a set of genes required for survival, representing several of the functional pathways identified, and used oligonucleotide-mediated recombineering (ORBIT)46 to disrupt their open reading frames. The initial genes selected were pafA (proteasome pathway), katG (catalase-peroxidase), recR (DNA repair), blaR (β-lactam sensing), and MAB_1456c (cobalamin synthesis). To control for non-specific effects of antibiotic selection during the recombineering process, we created a control strain using ORBIT to target a non-coding intergenic region downstream of a redundant tRNA gene (MAB_t5030c). We then individually screened each of these mutants to determine if they displayed deficits in survival by exposing cells to TIG/LZD, either in rich 7H9 media or under starvation conditions, as had been done in the Tn-Seq screen. For ΔkatG we detected clear defects in survival as soon as 3 days after antibiotic exposure in either rich media or under starvation conditions, corroborating the results of our Tn-Seq analysis (Figure 3A). We observed similar, albeit smaller, defects in the ΔpafA, ΔMAB_1456c, and ΔblaR, mutants under the conditions predicted by the screen, whereas we saw no defect in the ΔrecR mutant (Figure 3B-E). We also noted that several mutants displayed delays in growth recovery after antibiotic removal at times before survival defects were noted (Figure 3 Supplement). To further confirm the role of katG and pafA, and exclude off-target effects of recombineering, we performed genetic complementation analysis by restoring a wild-type copy of each gene into the respective ΔpafA and ΔkatG mutants. In each case, we integrated a single copy of the wild-type gene, under the control of its endogenous promoter, into the genome at the L5 attB site (hereafter pafA+, katG+ strains), and constructed isogenic control strains with an empty vector integrated at the same site (hereafter pafA-, katG-strains). We confirmed expression of the re-introduced copy of each gene by RT-qPCR in the pafA+, and katG+ strains, and found expression within roughly 2-fold of endogenous wild-type levels (Figure 4A,D). We then challenged these strains with TIG/LZD as before. In rich media, where the ΔkatG mutants have a moderate survival defect, the katG+ strain had roughly a 50-fold increase in viable cells relative to the katG-strain. We then exposed cells to antibiotics under starvation conditions, where the ΔkatG mutant phenotype is more severe. Under these conditions the katG-cells succumbed rapidly between 3d and 10d after antibiotic exposure, with a 1,000-fold decrease in viable cells relative to control cells, whereas the katG+ strain showed a near-complete restoration of antibiotic tolerance (Figure 4B). We analogously examined complementation of ΔpafA mutants, and although the phenotype of the ΔpafA mutant is less severe overall than a ΔkatG mutant, we saw a similar restoration of survival in pafA+ cells relative to pafA-cells (Figure 4E). We next evaluated whether the pafA-, and katG-strains were overall more sensitive to the growth inhibitory effects of TIG/LZD, or whether they had specific defects in survival above the mean bactericidal concentration. We performed MIC determination for TIG and LZD individually for each strain, comparing the katG+/katG- and pafA+/pafA- strains. We found that that the MICs for each of these strains were unchanged, demonstrating that these mutants were not more readily inhibited by these antibiotics. Instead, they have more rapid kinetics of cell death at bactericidal concentrations, consistent with a specific defect in antibiotic tolerance, and supporting a model whereby an initial growth-arresting inhibition of the direct antibiotic target can be uncoupled mechanistically from a distinct cell-death step (Figure 4C,F), as has been seen with other antibiotic classes26.

Revised version: To independently confirm a role in antibiotic tolerance for a set of genes from diverse pathways that were identified by Tn-Seq we selected a set of genes required for survival, representing several of the functional pathways identified, and used oligonucleotide-mediated recombineering (ORBIT)47 to disrupt their open reading frames. The initial genes selected were pafA (proteasome pathway), katG (catalase-peroxidase), recR (DNA repair), blaR (β-lactam sensing), and MAB_1456c (cobalamin synthesis). To control for non-specific effects of antibiotic selection during the recombineering process, we created a control strain using ORBIT to target a non-coding intergenic region downstream of a redundant tRNA gene (MAB_t5030c). We then individually screened each of these mutants to determine if they displayed deficits in survival by exposing cells to TIG/LZD, either in rich 7H9 media or under starvation conditions, as had been done in the Tn-Seq screen. For four out of five mutants we observed defects concordant with the Tn-Seq findings. For ΔkatG we detected clear defects in survival as soon as 3 days after antibiotic exposure in either rich media or under starvation conditions, corroborating the results of our Tn-Seq analysis (Figure 3A). We observed similar, albeit smaller, defects in the ΔpafA, ΔMAB_1456c, and ΔblaR, mutants under the conditions predicted by the screen (Figure 3B-D). We saw no survival defect in the ΔrecR mutant (Figure 3E). However, we observed that the ΔrecR mutant, as well as the ΔpafA mutant, instead displayed a marked defect in resumption of growth after removal of antibiotics (Figure 3 – figure supplement 1).

Validation of Tn-Seq results.

Earlier version: (A-E) Homologous recombination was used to delete the indicated genes, or to generate a control strain targeting a distant intergenic region distal to the non-essential tRNA gene MAB_t5030c. Each strain was either grown in 7H9 rich media or starved in PBS for 48 prior to the addition of antibiotics as indicated. The conditions tested here correspond to the conditions in the Tn-Seq analysis where a phenotype was observed. Data from additional conditions are in Figure 3 Supplemental. Comparisons in panels A-C are made to the same control strain but plotted independently for clarity. Error bars represent SEM, statistical significance is calculated at each time point using student’s t test. ****: p<0.0001, ***: p<0.001, **: p<0.01, *: p<0.05, ns: p>0.05. Antibiotics were added as described above. Data are representative of 4 independent experiments.

Revised version: (A-E) ORBIT homologous recombination was used to delete the indicated genes, or to generate a control strain targeting a distant intergenic region distal to the non-essential tRNA gene MAB_t5030c. Each strain was either grown in 7H9 rich media or starved in PBS for 48 prior to the addition of antibiotics as indicated. The conditions tested here correspond to the conditions in the Tn-Seq analysis where a phenotype was observed. Comparisons in panels A-C are made to the same control strain but plotted independently for clarity. Error bars represent SEM, statistical significance is calculated at each time point using student’s t test. ****: p<0.0001, ***: p<0.001, **: p<0.01, *: p<0.05, ns: p>0.05. Antibiotics were added as described above. Data are representative of 4 independent experiments.

Added in the revised version: To further confirm the role of katG and pafA, and exclude off-target effects of recombineering, we performed genetic complementation analysis by restoring a wild-type copy of each gene into the respective ΔpafA and ΔkatG mutants. In each case, we integrated a single copy of the wild-type gene, under the control of its endogenous promoter, into the genome at the L5 attB site (hereafter pafA+, katG+ strains), and constructed isogenic control strains with an empty vector integrated at the same site (hereafter pafA-, katG- strains). We confirmed expression of the re-introduced copy of each gene by RT-qPCR in the pafA+, and katG+ strains, and found expression within roughly 2-fold of endogenous wild-type levels (Figure 4A,D). We then challenged these strains with TIG/LZD as before. In rich media, where the ΔkatG mutants have a moderate survival defect, the katG+ strain had roughly a 50-fold increase in viable cells relative to the katG- strain. We then exposed cells to antibiotics under starvation conditions, where the ΔkatG mutant phenotype is more severe. Under these conditions the katG- cells succumbed rapidly between 3d and 10d after antibiotic exposure, with a 1,000-fold decrease in viable cells relative to control cells, whereas the katG+ strain showed a near-complete restoration of antibiotic tolerance (Figure 4B). We analogously examined complementation of ΔpafA mutants, and although the phenotype of the ΔpafA mutant is less severe overall than a ΔkatG mutant, we saw a similar restoration of survival in pafA+ cells relative to pafA- cells (Figure 4E). We next evaluated whether the pafA-, and katG-strains were overall more sensitive to the growth inhibitory effects of TIG/LZD, or whether they had specific defects in survival above the mean bactericidal concentration. We performed MIC determination for TIG and LZD individually for each strain, comparing the katG+/katG- and pafA+/pafA- strains. We found that the MICs for each of these strains were unchanged, demonstrating that these mutants were not more readily inhibited by these antibiotics (Figure 4C,F). Instead, they have more rapid kinetics of cell death at bactericidal concentrations, consistent with a specific defect in survival, and supporting a model whereby an initial growth-arresting inhibition of the direct antibiotic target can be uncoupled mechanistically from a distinct cell-death step, as has been seen with other antibiotic classes28.

Complementation analysis of katG and pafA mutants confirms their role in antibiotic tolerance.

(A) RT-qPCR analysis of katG expression in katG- (ΔkatG::pmv306), katG+ (ΔkatG::pmv306 katG), and control strain (ORBIT intergenic::pmv306). (B) CFU over time for katG+/katG- strains. (C) MICs for katG+/katG- strains. (D) Expression of pafA in pafA-(ΔpafA::pmv306), pafA+ (pafA::pmv306 pafA) and control strain. (E) CFU over time for pafA+/pafA- strains. (F) MICs for pafA+/pafA- strains. Antibiotic concentrations in (A-B, D-E) are as described above. Error bars represent SEM, statistical significance is calculated at each time point using student’s t test between katG+/katG- strains in (B) and between pafA+/pafA- strains in (E). ****: p<0.0001, ***: p<0.001, **: p<0.01, *: p<0.05, ns: p>0.05. Antibiotics were added as described above.

Reactive oxygen contributes to antibiotic lethality in Mabs

Earlier version: We next investigated the role of KatG and reactive oxygen in antibiotic tolerance more broadly. We began by assessing whether KatG conferred protection from other antibiotics with diverse mechanisms of action, selecting antibiotics that are used clinically for mycobacterial infections. Because katG- mutants showed the greatest defects in starvation-induced tolerance, we analyzed survival of katG+ and katG- strains in starvation-adapted cultures exposed to a panel of different antibiotics. Because both TIG and LZD both act by inhibiting translation, we began by exposing cells to either TIG or LZD alone. As expected, the degree of bacterial killing was significantly less with either agent alone than when they are added in combination. Upon exposure to either of these antibiotics the katG- cells died more rapidly than katG+ cells, though we note that the final proportion of persister cells in the population was unchanged in katG- cells (Figure 5A). When we exposed cells to rifabutin, an RNA polymerase inhibitor, we saw a similar effect, with a 100-fold loss of viability in katG- cells relative to the katG+ cells (Figure 5B). In contrast, when we exposed cultures to either levofloxacin (topoisomerase inhibitor) or cefoxitin (β-lactam inhibitor of peptidoglycan cross-linking), katG had little to no effect on cell viability (Figure 5C-D). Thus, the role of KatG is context-dependent, suggesting that some antibiotics generate oxidative stress that can be ameliorated by KatG while others do not.

Revised version: We next investigated the role of KatG and reactive oxygen in antibiotic tolerance more broadly. We began by assessing whether KatG conferred protection from other antibiotics with diverse mechanisms of action, selecting antibiotics that are used clinically for mycobacterial infections. Because katG- mutants showed the greatest defects in starvation-induced tolerance, we analyzed survival of katG+ and katG- strains in starvation-adapted cultures exposed to a panel of different antibiotics. Because both TIG and LZD both act by inhibiting translation, we began by exposing cells to either TIG or LZD alone. As expected, the degree of bacterial killing was significantly less with either agent alone than when they are added in combination. Upon exposure to either of these antibiotics the katG- cells died more rapidly than katG+ cells, though the final proportion of persister cells in the population was unchanged in katG- cells (Figure 5A). When we exposed cells to rifabutin, an RNA polymerase inhibitor, we saw a similar effect, with a 100-fold loss of viability in katG- cells relative to the katG+ cells (Figure 5B). In contrast, when we exposed cultures to either levofloxacin (topoisomerase inhibitor) or cefoxitin (β-lactam inhibitor of peptidoglycan cross-linking), katG had little to no effect on cell viability (Figure 5C-D). Thus, the role of KatG is context-dependent, suggesting that, in Mabs, some antibiotics generate oxidative stress that is ameliorated by KatG while others do not.

ROS-mediated toxicity following antibiotic exposure.

(A-D) Analysis of katG+/katG-cells challenged with different antibiotics. Cells were starved in PBS for 48h and then exposed to the indicated antibiotic. (E) Flow cytometry of control cells exposed to TIG/LZD (4h in 7H9 media or 72h in PBS) and then stained with DAPI and the ROS-sensitive dye cellROX green; percentage cellROX-positive cells are indicated. (F) Survival over time for aerated and hypoxic cultures of Mabs after exposure to TIG/LZD. (G) Survival over time for bipyridyl treated cells after exposure to TIG/LZD. Error bars represent SEM, statistical significance is calculated at each time point using student’s t test. ****: p<0.0001, ***: p<0.001, **: p<0.01, *: p<0.05, ns: p>0.05. (A-D, F) display combined data from 3 independent experiments. (E, G) are representative data from 3 independent experiments.

Earlier version: The identification of katG as essential for cells to survive exposure to TIG/LZD suggests that ROS are present and causing damage. Although TIG/LZD are translation inhibitors that do not directly generate ROS we evaluated whether they might nonetheless be triggering ROS accumulation as a secondary effect. We examined ROS levels in control Mabs using the ROS indicator dye cellROX, that is retained in cells when it becomes oxidized47. At baseline, during log-phase growth in rich media < 2% of cells had ROS accumulation (Figure 5E). We saw a moderate increase in ROS accumulation in starved cultures, with roughly 7% of the population cellROX+. However, when cells were exposed to antibiotics we saw a dramatic accumulation of ROS, with 57% of cells becoming cellROX+ when exposed to TIG/LZD in rich media and 33% of PBS-starved cells becoming cellROX+ when exposed to TIG/LZD. Taken together, these data indicate that translation inhibition does indeed have important downstream effects on cellular redox balance, with ROS accumulation that could be contributing to the lethal effects of antibiotics.

Revised version: The identification of katG as essential for cells to survive exposure to TIG/LZD suggests that ROS are present and causing damage. Although TIG and LZD are translation inhibitors that do not directly generate ROS we evaluated whether they might nonetheless be triggering ROS accumulation as a secondary effect. We examined ROS levels in control Mabs using the ROS indicator dye cellROX, that is retained in cells when it becomes oxidized48. At baseline, during log-phase growth in rich media < 2% of cells had ROS accumulation (Figure 5E). We saw a moderate increase in ROS accumulation in starved cultures, with roughly 7% of the population cellROX+. However, when cells were exposed to antibiotics we saw a dramatic accumulation of ROS, with 57% of cells becoming cellROX+ when exposed to TIG/LZD in rich media and 33% of PBS-starved cells becoming cellROX+ when exposed to TIG/LZD. Taken together, these data indicate that translation inhibition does indeed have important downstream effects on cellular redox balance, with ROS accumulation that could be contributing to the lethal effects of antibiotics.

Earlier version: We next tested whether ROS were contributing to cell death by reducing ROS production and assessing the impact on cell viability. A well-established system for studying hypoxia in mycobacteria is the Wayne Model of gradual-onset hypoxia, whereby low-density cultures are inoculated in sealed vessels with minimal headspace. As the culture slowly grows, the soluble oxygen is consumed, resulting in the slow onset of hypoxia over several days, a process that can be monitored by the decolorization of Methylene Blue dye in the media48. Under aerobic conditions in rich media, we observed the expected rapid killing of Mabs over the first 5 days with the combination of TIG/LZD, with more rapid loss of viability in KatG- cells. However, under hypoxic conditions, where ROS production is suppressed, we saw much slower bacterial killing. Importantly, under hypoxic conditions katG- cells no longer had a survival defect relative to katG+ cells, supporting the hypothesis that translation-inhibiting antibiotics also cause secondary accumulation of lethal ROS in antibiotic-treated cells that need to be detoxified by KatG (Figure 5F).

Revised version: We next tested whether ROS were contributing to cell death by reducing ROS production, and then assessing the impact on cell viability. A well-established system for studying hypoxia in mycobacteria is the Wayne Model of gradual-onset hypoxia, whereby low density cultures are inoculated in sealed vessels with minimal headspace. As the culture slowly grows, the soluble oxygen is consumed, resulting in the slow onset of hypoxia over several days, a process that can be monitored by the decolorization of Methylene Blue dye in the media49. Under aerobic conditions in rich media, we observed the expected rapid killing of Mabs over the first 5 days with the combination of TIG/LZD, with more rapid loss of viability in KatG- cells. However, under hypoxic conditions, where ROS production is suppressed, we saw much slower bacterial killing.

Added in the revised version: Importantly, under hypoxic conditions katG- cells no longer had a survival defect relative to katG+ cells, supporting the hypothesis that translation-inhibiting antibiotics also cause secondary accumulation of lethal ROS in antibiotic-treated cells that need to be detoxified by KatG (Figure 5F).

Earlier version: We also evaluated whether other methods of alleviating ROS damage might enhance survival. The iron chelator 2,2’-dipyridyl has been shown in other contexts to reduce ROS-mediated damage by suppressing the reaction of H202 with Fe2+ that generates highly oxidizing hydroxyl radicals (Fenton reaction), and which has been shown to mitigate oxidative damage in other bacteria following antibiotic exposure23, 26, 49. As seen in other bacteria, we find that in Mabs, 2,2’- dipyridy does indeed improve bacterial survival following exposure to bactericidal translation inhibitors, further supporting a role in ROS in cell death following translation inhibition (Figure 5G). We also assessed whether free radical scavengers like thiourea and 2,2,6,6- tetramethylpiperidine-1-oxyl (TEMPO) ameliorated antibiotic toxicity, although similar antioxidants had previously been shown in Mtb to increase respiration and paradoxically to increase ROS generation27. When we treated Mabs simultaneously with TIG/LZD in combination with either thiourea or TEMPO we did not observe a restoration of antibiotic tolerance, and, similar to observations in Mtb, actually observed increased bacterial cell death (Figure 5 Supplement).

Revised version: We also evaluated whether other methods of alleviating ROS damage might enhance survival. The iron chelator 2,2’-dipyridyl has been shown in other contexts to reduce ROS-mediated damage by suppressing the reaction of H202 with Fe2+ that generates highly oxidizing hydroxyl radicals (Fenton reaction), and which has been shown to mitigate oxidative damage in other bacteria following antibiotic exposure25, 28, 50. As seen in other bacteria, we find that in Mabs, 2,2’-dipyridy does indeed improve bacterial survival following exposure to bactericidal translation inhibitors, further supporting a role in ROS in cell death following translation inhibition (Figure 5G).

Added in the revised version: We also assessed whether free radical scavengers like thiourea and 2,2,6,6-tetramethylpiperidine-1-oxyl (TEMPO) ameliorated antibiotic toxicity, although similar thiol antioxidants had previously been shown in Mtb to increase respiration and ROS generation and to paradoxically decrease bacterial survival upon INH exposure30. When we treated Mabs simultaneously with TIG/LZD in combination with either thiourea or TEMPO we did not observe a restoration of antibiotic tolerance, and, similar to observations in Mtb, actually observed increased bacterial cell death (Figure 5 – figure supplement 1).

Removed in the revised version: Incomplete penetrance of katG phenotype among Mabs strains

Removed in the revised version: To test whether ROS accumulation was an effect occurring more broadly across different Mabs strains we obtained two clinical strains, exposed them to TIG/LZD in 7H9 media, and measured ROS accumulation with cellROX as above. We found that similar to ATCC 19977, both clinical Mabs strains had elevated ROS levels following translation inhibition (Figure 6A), suggesting that this is a conserved process in Mabs. Next, we tested the role of KatG in these Mabs clinical strains. We used ORBIT to disrupt the katG locus and evaluated the ability of these ΔkatG clinical strains to survive exposure to TIG/LZD under both stressed and un-stressed conditions. Unlike ROS accumulation, where the responses across strains were consistent, we saw a variable dependency on katG. Clinical strain-1 behaved differently overall, with no appreciable starvation-induced antibiotic tolerance, and no contribution of katG to survival. In contrast, for clinical strain-2, katG contributed significantly to starvation-induced antibiotic tolerance, behaving similarly to the ATCC 19977 reference strain. However, unlike the reference strain, katG was not required for survival of clinical strain-2 when exposed to antibiotics in 7H9 media (Figure 6B). Thus, although antibiotic-induced ROS accumulation was observed across all three Mabs strains, the ΔkatG phenotype displays incomplete penetrance, suggesting that in some Mabs strains alternative pathways exist that are able to compensate for the loss of katG.

Earlier version: Consistent antibiotic-induced ROS production but variable protection by KatG among different Mabs strains.

Revised version: Antibiotic-induced ROS accumulation is conserved but reliance on KatG is variable among Mabs strains

Added in the revised version: To test whether ROS accumulation was an effect occurring more broadly across different Mabs strains we obtained two clinical strains, exposed them to TIG/LZD in 7H9 media, and measured ROS accumulation with cellROX as above. We found that similar to ATCC 19977, both clinical

Added in the revised version: Mabs strains had elevated ROS levels following translation inhibition (Figure 6A), suggesting that this is a conserved process in Mabs. Next, we tested the role of KatG in these Mabs clinical strains. We used ORBIT to disrupt the katG locus and evaluated the ability of these ΔkatG clinical strains to survive exposure to TIG/LZD under both stressed and un-stressed conditions. Unlike ROS accumulation, where the responses across strains were consistent, we saw a variable dependency on katG. Clinical strain-1 behaved differently overall, with no appreciable starvation-induced antibiotic tolerance, and no contribution of katG to survival. In contrast, for clinical strain-2, katG contributed significantly to starvation-induced antibiotic tolerance, behaving similarly to the ATCC 19977 reference strain. However, unlike the reference strain, katG was not required for survival of clinical strain-2 when exposed to antibiotics in 7H9 media (Figure 6B). Thus, although antibiotic-induced ROS accumulation was observed across all three Mabs strains, the ΔkatG phenotype displays incomplete penetrance, suggesting that in some Mabs strains alternative pathways exist that are able to compensate for the loss of katG.

Added in the revised version: Conserved antibiotic-induced ROS production but variable protection by KatG among different Mabs strains.

(A) The indicated strains of Mabs were cultured in 7H9 media, exposed to TIG/LZD for 4h and then analyzed by cellROX staining. (B) CFU over time following TIG/LZD exposure. Error bars represent SEM, statistical significance is calculated at each time point using student’s t test. ****: p<0.0001, ***: p<0.001, **: p<0.01, *: p<0.05, ns: p>0.05. Combined data from 4 independent experiments are shown for persister survival experiments. Representative data from 2 independent experiments are shown for flow cytometry experiments. Antibiotics were added as described above.

Removed in the revised version: In summary, the results of these studies point to an important effect of ROS in amplifying the lethality of transcription and translation-inhibiting antibiotics in Mabs. Through genetic analysis we identified a number of ROS detoxification factors, including KatG, as necessary for survival in this context. This suggested that antibiotics induced an oxidative state in cells, and direct measurement of ROS following antibiotic exposure indicated that this was indeed the case. Further supporting the toxic effects of ROS in this context, we found that removal of oxygen both slowed bacterial killing and rendered KatG dispensable. Taken together, these results suggest that in Mabs antibiotic lethality is accelerated by ROS accumulation, and that survival requires active detoxification systems.

Discussion

Added in the revised version: The results of these studies point to an important effect of ROS in amplifying the lethality of transcription and translation-inhibiting antibiotics in Mabs. Through genetic analysis we identified a number of ROS detoxification factors, including KatG, as necessary for survival in this context. This suggested that antibiotics induced an oxidative state in cells, and direct measurement of ROS following antibiotic exposure indicated that this was indeed the case. Further supporting the toxic effects of ROS in this context, we found that removal of oxygen both slowed bacterial killing and rendered KatG dispensable. Taken together, these results suggest that in Mabs antibiotic lethality is accelerated by ROS accumulation, and that survival requires active detoxification systems.

Pathways necessary for antibiotic tolerance in Mabs

Earlier version: The phenomenon of antibiotic tolerance has been recognized for decades and has been observed in a broad array of bacterial species, but without a uniform underlying mechanism - suggesting that different pathways may play roles in different physiologic contexts. For example, the pathways identified in one bacterial species may not contribute to tolerance in another. In addition to E. coli, relA plays a role in tolerance in Pseudomonas aeruginosa53, Staphylococcus aureus54, and Mtb55. However, its role is not universal. Deletion of relA had no effect on antibiotic tolerance in Msmeg56, and in our Tn-Seq analysis, Mabs relA Tn mutants had no survival defect. However, a prior study of the Mabs relA mutant demonstrated that this strain still synthesizes (p)ppGpp57, suggesting potential genetic redundancy. Even within a single species there can be differences in the critical survival mechanisms depending on the context. In E. coli, RelA contributes to persister viability following exposure to β-lactams, but not quinolones53, and in our study we find KatG to be essential for tolerance to transcription and translation inhibitors but not to β-lactams or quinolones.

Revised version: The phenomenon of antibiotic tolerance has been recognized for decades and has been observed in a broad array of bacterial species10, 51–57 but without identification of a singular underlying mechanism conserved among species - suggesting that different pathways may play roles in different physiologic contexts. For example, the pathways identified in one bacterial species may not contribute to tolerance in another. In E. coli, relA plays an important role in stress-induced tolerance, as it does in Pseudomonas aeruginosa52, Staphylococcus aureus53, and Mtb54. However, its role is not universal. Deletion of relA had no effect on antibiotic tolerance in Msmeg55, and in our Tn-Seq analysis, Mabs relA Tn mutants had no survival defect. In the case of Mabs this may be due to genetic redundancy, as a prior study of the Mabs relA mutant demonstrated that this strain still synthesizes (p)ppGpp58. In addition, even within a single species there can be differences in the critical survival mechanisms depending on the context. In E. coli, RelA contributes strongly to persister viability following exposure to β-lactams, but not aminoglycosides56, and in our study we find KatG to be essential for tolerance to transcription and translation inhibitors but not to a β-lactam (cefoxitin) or a quinolone (levofloxacin).

Mechanisms of antibiotic lethality

Earlier version: Our findings strongly support the idea that ROS can be a significant contributor to antibiotic lethality in Mabs and contribute to the growing evidence that this is conserved across mycobacteria. In Mabs, Mtb and Msmeg, other groups have also observed reduced antibiotic-mediated killing in hypoxic conditions25, 37. In addition, in Mtb, exposure to rifampin also generates ROS19, 60, and katG contributes to survival in rifampin-treated cells26.

Revised version: Our findings strongly support the idea that antibiotic-induced ROS can be a significant contributor to bactericidal activity in Mabs and contribute to the growing evidence that this phenomenon is conserved across diverse types of bacteria. In mycobacteria, other groups have observed that hypoxia reduced antibiotic-mediated killing in Mabs, Mtb and Msmeg, and in Mtb, exposure to rifampin or moxifloxacin also generates ROS, with katG contributing to survival in rifampin-treated cells28, 29. Notably, we found evidence for ROS mediated bactericidal activity with the translation inhibiting antibiotics tigecycline and linezolid, as well as the transcription inhibiting antibiotic rifabutin. This contrasts with findings in E. coli where the rifamycins and tetracyclines are not bactericidal and do not induce ROS25, 26, suggesting that mycobacteria-specific responses may exist that result in lethal ROS production.

Earlier version: Exactly how transcription or translation blockade leads to increased ROS levels is also not known. In principle, any of several derangements could lead to ROS accumulation. One of the major sources of cellular ROS is oxidative phosphorylation, as hydrogen peroxide and superoxide are natural byproducts. Thus, increased ROS generation by oxidative phosphorylation is an attractive hypothesis. Alternatively, particularly under starvation conditions, it is possible that antioxidants and ROS scavengers may become depleted, creating a more oxidizing environment. Our Tn-Seq analysis provides additional insight on this. We noted a small class of Tn mutants that were paradoxically protected from antibiotic lethality (Figure 2B). Prominent among this class of mutants were several independent components of the NADH dehydrogenase complex. Also known as Complex I of the electron transport chain, it is one of the key entry points for electrons into the oxidative phosphorylation pathway. The observation that mutants lacking this complex are protected suggests that decreasing flux through oxidative phosphorylation, with a concomitant decrease in ROS generation, may enhance survival during antibiotic exposure. A mechanistic understanding of how blockade of either transcription or translation leads to deranged oxygen utilization is an unresolved question that will require further study.

Revised version: Exactly how transcription or translation blockade leads to increased ROS levels is not known. In principle, any of several derangements could lead to ROS accumulation. One of the major sources of cellular ROS is oxidative phosphorylation, as hydrogen peroxide and superoxide are natural byproducts. Thus, increased ROS generation by oxidative phosphorylation is an attractive hypothesis. Alternatively, particularly under starvation conditions, it is possible that antioxidants and ROS scavengers may become depleted, creating a more oxidizing environment. Our Tn-Seq analysis provides additional insight on this. We noted a small class of Tn mutants that were paradoxically protected from antibiotic lethality (Figure 2B). Prominent among this class of mutants were several independent components of the NADH dehydrogenase complex. Also known as Complex I of the electron transport chain, it is one of the key entry points for electrons into the oxidative phosphorylation pathway. The observation that mutants lacking this complex are protected suggests that decreasing flux through oxidative phosphorylation, with a concomitant decrease in ROS generation, may enhance survival during antibiotic exposure. A mechanistic understanding of how blockade of either transcription or translation leads to deranged oxygen utilization is an unresolved question that will require further study.

Earlier version: ROS accumulation is not universal following exposure to bactericidal antibiotics. While antibiotic-induced ROS is well-documented, under some conditions, such as higher concentrations of antibiotic, multiple studies have also observed antibiotic lethality without ROS accumulation23, 58–61. Similarly, prior studies have found that under certain conditions, E. coli mutants lacking catalase have defects in antibiotic tolerance62–64, whereas under other conditions they do not23, 65. In Mabs the role of katG was also not uniform, as it had no impact on survival following exposure to levofloxacin or cefoxitin. Additional studies will be needed to determine whether levofloxacin and cefoxitin kill without generating ROS, or whether ROS is generated but effectively detoxified in the absence of katG. This latter possibility is suggested by our findings in Mabs clinical strains. While we saw ROS accumulation in both strains after TIG/LZD exposure, the role of katG was variable between the strains. This suggests that compensatory pathways likely exist and that in Mabs they can overcome the loss of katG.

Revised version: ROS accumulation is not universal following exposure to bactericidal antibiotics. While antibiotic-induced ROS is well-documented, under some conditions, such as higher concentrations of antibiotic, multiple studies have also observed antibiotic lethality without ROS accumulation59–62. Similarly, prior studies have found that under certain conditions, E. coli mutants lacking catalase have defects in antibiotic tolerance63–65, whereas under other conditions they do not59 . In Mabs the role of katG was also not uniform, as it had no impact on survival following exposure to levofloxacin or cefoxitin. Additional studies will be needed to determine whether levofloxacin and cefoxitin kill without generating ROS, or whether ROS is generated but effectively detoxified by other systems in the absence of katG. This latter possibility is suggested by our findings in Mabs clinical strains. While we saw ROS accumulation in all strains after TIG/LZD exposure, the role of katG was variable between the strains. This suggests that compensatory pathways likely exist and that in Mabs they can overcome the loss of katG.

Therapeutic implications

Earlier version: Mabs infections are particularly challenging to treat, with relapse rates over 50%3. Our results highlight several bacterial processes, such as the bacterial proteasome and ROS detoxification, which might be targeted therapeutically to reduce the survival of antibiotic tolerant bacteria in patients with Mabs infection. Agents targeting these processes might not have any intrinsic antimicrobial activity alone, but might act to disrupt the unique physiology required to survive antibiotics. This would represent a new therapeutic class of “persistence inhibitors” that might act synergistically with traditional antibiotics to eliminate the subpopulation of cells that would otherwise remain viable despite prolonged antibiotic treatment in patients with Mabs and other chronic infections6, 27.

Revised version: Mabs infections are particularly challenging to treat, and frequently have poor outcomes4. Our results highlight several bacterial processes, such as the bacterial proteasome and ROS detoxification, which might be targeted therapeutically to reduce the survival of antibiotic tolerant bacteria in patients with Mabs infection. Agents targeting these processes might not have any intrinsic antimicrobial activity alone, but might act to disrupt the unique physiology required to survive antibiotics. This would represent a new therapeutic class of “persistence inhibitors” that might act synergistically with traditional antibiotics to eliminate the subpopulation of cells that would otherwise remain viable, despite prolonged antibiotic treatment, in patients with Mabs and other chronic infections28, 30.

Limitations

Earlier version: Tn-Seq has inherent drawbacks, including an inability to identify mutants in essential genes, or in cases of genetic redundancy. Thus, there are likely genes needed for antibiotic tolerance in Mabs that were not identified in this study. In addition, we studied the response to a single class of antibiotic, focusing on the translation inhibitors often used to treat Mabs infections, and we studied only spontaneous persister cells and starvation-induced antibiotic tolerance. It is likely that examining other antibiotics, with different mechanisms of action, or different stresses that induce tolerance would identify additional genes contributing to survival and would allow identification of core pathways that might be shared in differing physiologic contexts of antibiotic and stress25, 27.

Revised version: Tn-Seq has inherent drawbacks, including an inability to identify mutants in essential genes, or in cases of genetic redundancy. Thus, there are likely genes needed for antibiotic tolerance in Mabs that were not identified in this study. In addition, we studied the response to a single class of antibiotic, focusing on the translation inhibitors often used to treat Mabs infections, and we studied only spontaneous persister cells and starvation-induced antibiotic tolerance. It is likely that examining other antibiotics, with different mechanisms of action, or different stresses that induce tolerance would identify additional genes contributing to survival and would allow identification of core pathways that might be shared in differing physiologic contexts of antibiotic and stress.

Materials and methods

Key Reagent Table

Bacterial strains and culture conditions

Earlier version: Mabs ATCC 19977, clinical Mabs strains and Msmeg (MC2 155) were grown in BD Middlebrook 7H9 media (liquid) or 7H10 media (solid) supplemented with 0.5% glycerol (Sigma) and 0.2% Tween-80 (Fisher) but without any OADC supplementation except for transformations. Sacramento clinical isolates were obtained from the Sacramento County Department of Public Health Mycobacteriology Laboratory. Confirmation of clinical isolates as Mabs was performed by amplifying the 16s rRNA locus and Sanger sequencing (maybe reference new data table?). Mtb (Erdman) was grown in 7H9 (liquid) or 7H10 (solid) supplemented with 0.5% glycerol, 0.1% Tween-80, and 10% OADC (BD). All cultures were grown at 37°C with gentle shaking. Except for specific hypoxia conditions, all liquid cultures were grown with 90% container headspace or using a gas permeable cap to ensure culture oxygenation. PBS starvation was achieved by washing OD0.5-1.0 Mabs 1X in DPBS (-Ca/Mg, Gibco) and resuspending in DPBS at OD=1, supplemented with 0.1% tyloxapol (Sigma).

Revised version: Mabs ATCC 19977, clinical Mabs strains and Msmeg (MC2 155) were grown in BD Middlebrook 7H9 media (liquid) or 7H10 media (solid) supplemented with 0.5% glycerol (Sigma) and 0.2% Tween-80 (Fisher) but without any OADC supplementation except for transformations. Sacramento clinical isolates were obtained from the Sacramento County Department of Public Health Mycobacteriology Laboratory. Confirmation of clinical isolates as Mabs was performed by amplifying the 16s rRNA locus and Sanger sequencing (Supplemental Table 4). Mtb (Erdman) was grown in 7H9 (liquid) or 7H10 (solid) supplemented with 0.5% glycerol, 0.1% Tween-80, and 10% OADC (BD). All cultures were grown at 37°C with gentle shaking. Except for specific hypoxia conditions, all liquid cultures were grown with 90% container headspace or using a gas permeable cap to ensure culture oxygenation. PBS starvation was achieved by washing OD 0.5-1.0 Mabs 1X in DPBS (-Ca/Mg, Gibco) and resuspending in DPBS at OD=1, supplemented with 0.1% tyloxapol (Sigma).

Mabs antibiotic experiments

Earlier version: For PBS starvation experiments, stocks of Mabs were grown for 48 hours in 7H9 passaging continuously in log phase, then PBS starved or passaged in log phase for an additional 48h. Log phase or PBS starved Mabs were then resuspended in antibiotic containing media at OD1.0. For experiments with hypoxia, Mabs in mid-log aerobic growth was adjusted to OD0.001 in media with 1.5μg/ml methylene blue and added to a rubber septum sealed glass vial with 50% headspace. Methylene blue discoloration was observed at 3d and antibiotics were added at 5d. For thiourea and TEMPO experiments, these were added to the cultures at the time of antibiotic administration. For bipyridyl experiments, 62.5μM bipyridyl was added to the cultures 2h prior to antibiotic administration. We empirically determined the half-life of each antibiotic in 7H9 media at 37-deg and for those with half-lives shorter than the experiment, supplemented cultures to with additional antibiotic to maintain the concentration of active antibiotic. Antibiotics were used at the following concentrations: tigecycline (Chem-Impex) at 10μg/ml (8 fold above MIC, re-administered every 3 days), linezolid (Chem-Impex) at 100μg/ml (20 fold above MIC), levofloxacin (Sigma) at 40μg/ml (8 fold above MIC), cefoxitin (Chem Impex) at 80μg/ml (8 fold above MIC, re-administered every 3 days), and rifabutin (Cayman) at 40μg/ml (4 fold above MIC). After antibiotic administration, colony forming units over time were measured by plating on 7H10 solid media. For experiments where growth recovery time in liquid media was quantified, 100μl of sample was removed at d6 after antibiotic administration and washed 2X in antibiotic-free media. The samples were resuspended in 5ml of antibiotic free media, and OD620 measurements were taken with a FilterMax F3 plate reader (Molecular Devices) until maximum cell density (OD of ∼5.0) was reached.

Revised version: For PBS starvation experiments, stocks of Mabs were grown for 48 hours in 7H9 passaging continuously in log phase, then either PBS starved or passaged in log phase for an additional 48 h. Log phase or PBS starved Mabs were then resuspended in antibiotic containing media at OD1.0. For experiments with hypoxia, Mabs in mid-log aerobic growth was adjusted to OD 0.001 in media with 1.5 μg/ml methylene blue and added to a rubber septum sealed glass vial with 50% headspace. Methylene blue discoloration was observed at 3d and antibiotics were added at 5d. For thiourea and TEMPO experiments, these were added to the cultures at the time of antibiotic administration. For bipyridyl experiments, 62.5 μM bipyridyl was added to the cultures 2h prior to antibiotic administration. We empirically determined the half-life of each antibiotic in 7H9 media at 37-deg and for those with half-lives shorter than the experiment, supplemented cultures to with additional antibiotic to maintain the concentration of active antibiotic. Antibiotics were used at the following concentrations: tigecycline (Chem-Impex) at 10 μg/ml (8-fold above MIC, re-administered every 3 days), linezolid (Chem-Impex) at 100 μg/ml (20-fold above MIC), levofloxacin (Sigma) at 40μg/ml (8-fold above MIC), cefoxitin (Chem Impex) at 80 μg/ml (8-fold above MIC, re-administered every 3 days), and rifabutin (Cayman) at 40μg/ml (4-fold above MIC). After antibiotic administration, colony forming units over time were measured by plating on 7H10 solid media. For experiments where growth recovery time in liquid media was quantified, 100μl of sample was removed at d6 after antibiotic administration and washed 2X in antibiotic-free media. The samples were resuspended in 5ml of antibiotic free media, and OD620 measurements were taken with a FilterMax F3 plate reader (Molecular Devices) until maximum cell density (OD of ∼5.0) was reached.

Msmeg antibiotic experiments

Individual colonies were picked and grown for 48 hours in log phase before being PBS starved or passaged in log phase for 48h. Log phase or PBS starved Msmeg were then resuspended in antibiotic containing media at OD1.0. Antibiotics were used at the following concentrations: tigecycline (Chem-Impex) at 1.25 μg/ml (8-fold above MIC, re-administered every 3 days), linezolid (Chem-Impex) at 2.5 μg/ml (8 fold above MIC), rifampin (Sigma) at 32 μg/ml (8-fold above MIC, re-administered every 6 days), isoniazid (Sigma) at 32 μg/ml (8-fold above MIC, re-administered every 6 days), and ethambutol (Thermo) at 4 μg/ml (8-fold above MIC, re-administered every 3 days). After antibiotic administration, colony forming units over time were measured.

Mtb antibiotic experiments

Earlier version: Freezer stocks of Mtb were thawed and grown for 5-7 days in log phase before being starved for 14d or longer. Non-starved control Mtb were thawed such that they were also grown for 5-7 days in log phase before experimental use. Log phase or PBS starved Mtb was then resuspended in antibiotic containing media and adjusted to OD1.0. Antibiotics were used at the following concentrations: rifampin (Sigma) at 0.1μg/ml (4 fold above MIC, re-administered every 6 days), isoniazid (Sigma) at 0.1μg/ml (4 fold above MIC, re-administered every 6 days), and ethambutol (Thermo) at 8μg/ml (4 fold above MIC, re-administered every 6 days). After antibiotic administration, colony forming units over time were measured.

Revised version: Freezer stocks of Mtb were thawed and grown for 5-7 days in log phase before being starved for 14d or longer. Non-starved control Mtb were thawed such that they were also grown for 5-7 days in log phase before experimental use. Log phase or PBS starved Mtb was then resuspended in antibiotic containing media and adjusted to OD 1.0. Antibiotics were used at the following concentrations: rifampin (Sigma) at 0.1 μg/ml (4-fold above MIC, re-administered every 6 days), isoniazid (Sigma) at 0.1μg/ml (4-fold above MIC, re-administered every 6 days), and ethambutol (Thermo) at 8 μg/ml (4-fold above MIC, re-administered every 6 days). After antibiotic administration, colony forming units over time were measured.

Transposon insertion sequencing

Removed in the revised version: The construction of this Himar1 transposon Tn library has been described previously33. Screening was performed by growing a freezer stock of the library for 2.5 days in log phase before 48-hour PBS starvation or further continuous log-phase growth. Samples were then resuspended in media containing either tigecycline/linezolid or an equal volume of DMSO solvent and incubated for 6 days, with a re-administration of tigecycline or matching DMSO on day 3. Cultures were aerated by culturing in a vented cap bottle with gentle agitation at 40 revolutions per minute throughout the experiment. The samples were then washed 2X in antibiotic free liquid media, resuspended in antibiotic free liquid media (10X the original culture volume), and grown until OD0.5-1.0.

Earlier version: Subsequently, the samples underwent three more rounds of 100-fold passaging in liquid media to amplify surviving bacteria before the samples were collected in Trizol (Invitrogen). A sample taken at the time of the commencement of PBS starvation was collected in Trizol and used as the input control. Three independent trials of this experiment were submitted to the UC Davis DNA Technologies Core, where Tn insertion site flanking sequences were amplified as described previously33 and sequenced on an Element Biosciences AVITI. Sequence reads were mapped to the ATCC 19977 genome and analyzed using TRANSIT software with the following parameters: 0% of N/C termini ignored, 10,000 samples, TTR normalization, LOESS correction, include sites with all zeros, site restricted resampling. Genes with significant changes were defined as those with adjusted p-value (p-adj.) <0.05 and log2 fold change >0.5. P-adj. was calculated using the Benjamini-Hochberg correction.

Revised version: The construction of this Himar1 transposon Tn library has been described previously34. Screening was performed by growing a freezer stock of the library for 2.5 days in log phase before 48 h PBS starvation or further continuous log-phase growth. Samples were then resuspended in media containing either tigecycline/linezolid or an equal volume of DMSO solvent and incubated for 6 days, with a re-administration of tigecycline or matching DMSO on day 3. Cultures were aerated by culturing in a vented cap bottle with gentle agitation at 40 revolutions per minute throughout the experiment. The samples were then washed 2X in antibiotic free liquid media, resuspended in antibiotic free liquid media (10X the original culture volume), and grown until OD=0.5-1.0. Subsequently, the samples underwent three more rounds of 100-fold passaging in liquid media to amplify surviving bacteria before the samples were collected in Trizol (Invitrogen). A sample taken at the time of the commencement of PBS starvation was collected in Trizol and used as the input control. Three independent trials of this experiment were submitted to the UC Davis DNA Technologies Core, where Tn insertion site flanking sequences were amplified as described previously34 and sequenced on an Element Biosciences AVITI. Sequence reads were mapped to the ATCC 19977 genome and analyzed using TRANSIT software with the following parameters: 0% of N/C termini ignored, 10,000 samples, TTR normalization, LOESS correction, include sites with all zeros, site restricted resampling. Genes with significant changes were defined as those with adjusted p-value (p-adj.) <0.05 and log2 fold change >0.5. P-adj. was calculated using the Benjamini-Hochberg correction.

Pathway enrichment analysis

Earlier version: To improve gene annotation, Mabs orthologs to Mtb genes were identified. Mabs genes were first converted into protein sequences using Mycobrowser, and protein sequences were then used to perform reciprocal BLASTp searches. Mabs genes and Mtb genes that mapped to each other using independent one-way BLASTp searches with a maximum e-value cutoff of 0.1 were considered orthologs. For pathway analysis, gene lists (Mtb orthologs) were then imported into the DAVID knowledgebase42 and pathway enrichment analysis performed for Gene Ontology biological process, Uniprot keyword and KEGG databases with statistical analysis Fisher’s exact test and nominal p-value reported.

Revised version: To improve gene annotation, Mabs orthologs to Mtb genes were identified. Mabs genes were first converted into protein sequences using Mycobrowser, and protein sequences were then used to perform reciprocal BLASTp searches. Mabs genes and Mtb genes that mapped to each other using independent one-way BLASTp searches with a maximum e-value cutoff of 0.1 were considered orthologs. For pathway analysis, gene lists (Mtb orthologs) were then imported into the DAVID knowledgebase43 and pathway enrichment analysis performed for Gene Ontology biological process, Uniprot keyword and KEGG databases with statistical analysis Fisher’s exact test and nominal p-value reported.

Gene deletion and complementation

Earlier version: Knockout strains were generated using ORBITT48. Briefly, Mabs was transformed with the kanamycin-resistant ORBIT recombineering plasmid pkm444. 20ml Mabs at OD0.5-1.0 was washed 2X in 10% glycerol and resuspended in 200ul 10% glycerol. 500ng plasmid was added and electroporated at 2.5kV in 0.2cm cuvettes. The bacteria were allowed to recover overnight before plating on 150μg/ml kanamycin plates. Clones were selected and regrown in liquid media supplemented with 150μg/ml kanamycin and 10% OADC (BD) to OD0.5-1.0. For recombineering, the pkm444-Mabs was grown to mid-log, then diluted to OD0.1 and 200mM glycine (Fisher) was added to the media. 16 hours later, 500mM sucrose (Sigma) and 500ng/ml anhydrotetracycline (Cayman) were added and incubated for an additional 4 hours. Subsequently, the Mabs was washed 2X in ice cold 10% glycerol + 500mM sucrose. 200ul of 10X concentrated Mabs was then electroporated with 600ng of the zeocin-R ORBIT payload pkm496 plasmid and 2μg of targeting oligonucleotide (Table S3) at 2.5kV in 0.2cm cuvettes. The Mabs was then allowed to recover overnight in liquid media with 10% OADC and 500ng/ml anhydrotetracycline before being plated on 150μg/ml zeocin plates. Mutants were then selected and screened for gene deletion by PCR amplification and Sanger sequencing. For genetic complementation, the endogenous loci including promoter and terminator sequences were amplified by PCR and cloned into the EcoRV site of pmv306 with kanamycin resistance66. In the case of katG, the upstream gene furA was also included in the complementation construct to achieve optimal katG expression.

Revised version: Knockout strains were generated using ORBIT47. Briefly, Mabs was transformed with the kanamycin-resistant ORBIT recombineering plasmid pkm444. 20ml Mabs at OD=0.5-1.0 was washed 2X in 10% glycerol and resuspended in 200ul 10% glycerol. 500ng plasmid was added and electroporated at 2.5kV in 0.2cm cuvettes. The bacteria were allowed to recover overnight before plating on 150μg/ml kanamycin plates. Clones were selected and regrown in liquid media supplemented with 150μg/ml kanamycin and 10% OADC (BD) to OD0.5-1.0. For recombineering, the pkm444-Mabs was grown to mid-log, then diluted to OD0.1 and 200mM glycine (Fisher) was added to the media. 16 h later, 500mM sucrose (Sigma) and 500 ng/ml anhydrotetracycline (Cayman) were added and incubated for an additional 4 h. Subsequently, the Mabs was washed 2X in ice cold 10% glycerol + 500mM sucrose. 200ul of 10X concentrated Mabs was then electroporated with 600ng of the zeocin-R ORBIT payload pkm496 plasmid and 2μg of targeting oligonucleotide (Table S3) at 2.5kV in 0.2cm cuvettes. The Mabs was then allowed to recover overnight in liquid media with 10% OADC and 500ng/ml anhydrotetracycline before being plated on 150 μg/ml zeocin plates. Mutants were then selected and screened for gene deletion by PCR amplification and Sanger sequencing. For genetic complementation, the endogenous loci including promoter and terminator sequences were amplified by PCR and cloned into the EcoRV site of pmv306 with kanamycin resistance66. In the case of katG, the upstream gene furA was also included in the complementation construct to achieve optimal katG expression.

MIC determination

Earlier version: Two-fold serial dilutions of antibiotics were prepared in a 96 well plate in 100ul volume. 100ul of 2X bacteria were added (for Mabs: used a final OD of 0.001, Msmeg: OD0.001, Mtb: OD0.01), making a final volume of 200μl. The plates were incubated until there was visible growth in the no antibiotic control well. At this time, the bacteria were transferred to a new plate with 20μl of 40% paraformaldehyde and OD620 measurements were taken with a FilterMax F3 plate reader (Molecular Devices).

Revised version: Two-fold serial dilutions of antibiotics were prepared in a 96 well plate in 100ul volume. 100ul of 2X bacteria were added (for Mabs: used a final OD of 0.001, Msmeg: OD0.001, Mtb: OD0.01), making a final volume of 200μl. The plates were incubated until there was visible growth in the no antibiotic control well. At this time, the bacteria were transferred to a new plate with 20μl of 40% paraformaldehyde and OD620 measurements were taken with a FilterMax F3 plate reader (Molecular Devices). MIC values for wild-type ATCC 19977 under these conditions are listed in Supplemental Table 5.

Flow cytometry

Earlier version: OD1 Mab was stained with cellROX green (Invitrogen) at a final concentration of 5μM for 1hr at 37C. The cells were then washed in PBS and resuspended in PBS with 4% paraformaldehyde and 5μg/ml DAPI (Sigma). The samples were run on an LSRII flow cytometer (BD). Fluorophores were excited with the 405nm (DAPI) and 488nm (cellROX) lasers. Detection was performed using the 450/50 (505LP) filter for DAPI and a 525/50 (555LP) filter for cellROX. Data were analyzed with FlowJo software (BD).

Revised version: A culture of OD=1.0 Mabs was stained with cellROX green (Invitrogen) at a final concentration of 5μM for 1hr at 37C. The cells were then washed in PBS and resuspended in PBS with 4% paraformaldehyde and 5 μg/ml DAPI (Sigma). The samples were run on an LSRII flow cytometer (BD). Fluorophores were excited with the 405nm (DAPI) and 488nm (cellROX) lasers. Detection was performed using the 450/50 (505LP) filter for DAPI and a 525/50 (555LP) filter for cellROX. Data were analyzed with FlowJo software (BD).

DNA/RNA Purification

Samples were resuspended in 5 volumes of Trizol, and bead beat with 0.1mm zirconia beads (Biospec) 6x2min at 4°C in a Mini-Beadbeater-16 (Biospec). Chloroform was added and RNA in the aqueous phase removed. For DNA isolation, a second RNA extraction was performed with 0.8M guanidine thiocyanate and 0.5M guanidine hydrochloride, 60mM Acetate pH 5.2, 1mM EDTA. DNA was then isolated with back-extraction buffer (4M Guanidine Thiocyanate, 50mM Sodium Citrate, 1M Tris base (without pH adjustment ∼pH 11) and DNA purified using a PureLink RNA Mini Kit (Invitrogen).

RT-qPCR

RNA was purified using PureLink RNA Mini Kit per manufacturer’s instructions. The samples were DNAseI (NEB) treated for 15min/37°C before stopping the reaction by adding 3.5mM EDTA and heating for 10min/75°C. cDNA was synthesized from 500ng total RNA using random hexamers and Maxima H minus reverse transcriptase (Thermo). No reverse-transcription controls were also included and used to confirm the lack of genomic DNA-driven amplification. qPCR reactions used Taq polymerase (NEB) and EvaGreen (Biotium) and were run on Biorad CFX Opus 96 Real-Time PCR System. Melt curves were included for each sample to confirm uniform amplicon identity between samples. Gene-specific amplification was quantified by comparison to a standard curve generated from 3-fold serial dilutions of a control sample, then normalized to 16S rRNA within each sample.